---
# More about metadata section
title: "Learnr: and example"
author: "D. Puthier"
date: "`r date()`"
#output: learnr::tutorial
output:
    html_document:
    toc: true
runtime: shiny_prerendered
# Do not index/display tutorial by setting `private: true`
private: true
toc: true
description: >
  This is a demo tutorial.
---

\tableofcontents

# Installation

**WARNING:** Currently you need to install learnr from commit "bd62765" which correspond to PR #310 (from Zoë Wilkinson Saldaña). Any attempt to try installation from another branch is expected to fail. 

**Note:** This is just for demonstration pupose and bash/python support in learnr is still in development.

```{r learnr, include=TRUE, eval=FALSE}
devtools::install_github(repo="rstudio/learnr", ref = "bd62765", force = FALSE)
```

**WARNING:** Currently you need to install htmltools version 0.3.6. The current version is compatible with master branch of learnr but our version lacks a bug fix.

```{r htmltools, include=TRUE, eval=FALSE}
require(devtools)
install_version("htmltools", version = "0.3.6", repos = "http://cran.us.r-project.org")
```




# Starting an exercice

- First you need to import the learnr package. If you don't, exercice snippets won't turn into nice javascript form. 

- Also load *gradethis* to evaluate students outputs and set the exercice checker function to gradethis::grade_learnr.


```{r}
library(learnr)
library(gradethis)
tutorial_options(exercise.checker = gradethis::grade_learnr, exercise.completion = TRUE)
```


# Grade a learnr exercice 

## Using shell as engine


**WARNING:** due to knitr limitations each exercices is run in a new terminal As a consequence the working directory will be redefined each time a new snippet is created. One may perhaps have a look at [knitr hooks](:https://yihui.org/knitr/hooks/) to get a workaround.

**Exercice:** go to /tmp directory and call pwd command.

```{sh grade_sh, exercise.engine="sh", exercise=TRUE, exercise.cap = "Bash exo #1"}

```

```{sh grade_sh-solution}
cd /tmp
pwd
```

```{sh grade_sh-check}
grade_code("Good job !!!!")
```

## Using R engine

**Question:** Assign 10 to a variable named x and print(x)

```{r r_exo_1,  exercise=TRUE, exercise.cap = "R exo #1"}

```


```{r r_exo_1-solution}
x <- 10
print(x)
```

```{r r_exo_1-check}
grade_code("You are mastering R !")
```

## Using Python engine

Concatenate "Hello" and "world" and print the result

```{python python_exo_1,  exercise=TRUE, exercise.cap = "R exo #1"}

```


```{python python_exo_1-solution}
print("Hello" + "world")
```

```{python python_exo_1-check}
grade_code("Astonishing !")
```

# Session info

```{r}
sessionInfo()
```




